# Getting started

Topics:
1. How to print
2. How to assign variables
3. How to comment
4. Syntax for basic math

[Link to the course](https://juliaacademy.com/p/intro-to-julia)

[Link to the notebook](https://github.com/JuliaAcademy/Introduction-to-Julia)

## How to print

In Julia we usually use `println()` to print

```julia
println("I'm excited to learn Julia!")
```

## How to assign variables

All we need is a variable name, value, and an equal's sign!<br>
Julia will figure out types for us.

```julia
my_answer = 42
typeof(my_answer)
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

```julia
😺 = "smiley cat!"
typeof(😺)
```

To type a smiley cat, use tab completion to select the emoji name and then tab again

```julia
# \:smi + <tab> --> select with down arrow + <enter> ---> <tab> + <enter> to complete
```

After assigning a value to a variable, we can reassign a value of a different type to that variable without any issue.

```julia
😺 = 1
```

```julia
typeof(😺)
```

Note: Julia allows us to write super generic code, and 😺 is an example of this. 

This allows us to write code like

```julia
😀 = 0
😞 = -1
```

```julia
😺 + 😞 == 😀
```

## How to comment

```julia
# You can leave comments on a single line using the pound/hash key
```

```julia
#=

For multi-line comments, 
use the '#= =#' sequence.

=#
```

## Syntax for basic math

$$
    f(x) = 2
$$


```julia
sum = 3 + 7
```

```julia
difference = 10 - 3
```

```julia
product = 20 * 5
```

```julia
quotient = 100 / 10
```

```julia
power = 10 ^ 2
```

```julia
modulus = 101 % 2
```

### Exercises

#### 1.1 
Look up [docs](https://docs.julialang.org/en/v1/) for the `convert` function.

```julia
convert(Int64, 3.0)
```

#### 1.2 
Assign `365` to a variable named `days`. Convert `days` to a float and assign it to variable `days_float`

```julia
days = 365
days_float = convert(Float64, days)
```

`@assert`: Throw an AssertionError `if cond is false`. Preferred syntax for writing assertions. Message text is optionally displayed upon assertion failure.

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert iseven(3) "3 is an odd number!"

```

```julia
@assert isodd(3) "What even are numbers?"
```

#### 1.3 
See what happens when you execute

```julia
convert(Int64, "1")

parse(Int64, "1")
```

```julia
convert(Int64, "1")
```

This happens because you cannott `convert` an object of type `String` to an object of type `Int64`

```julia
parse(Int64, "1")
```